rule_files:
  - model-registry-operator-alerting.rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="model-registry-operator"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="model-registry-operator"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="model-registry-operator"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="model-registry-operator"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="model-registry-operator"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Model Registry Operator Probe Success Burn Rate
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="model-registry-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="model-registry-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: Model Registry Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Model Registry Operator Probe Success Burn Rate
              instance: "model-registry-operator"
              namespace: "redhat-ods-applications"
              severity: critical
            exp_annotations:
              summary: "Model Registry Operator Probe Success Burn Rate"
              message: "High error budget burn for model-registry-operator (current value: 3)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-model-registry-operator-probe-success-burn-rate.md"

  - interval: 1m
    input_series: 
      - series: probe_success:burnrate30m{instance="model-registry-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="model-registry-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: Model Registry Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Model Registry Operator Probe Success Burn Rate
              instance: "model-registry-operator"
              namespace: "redhat-ods-applications"
              severity: critical
            exp_annotations:
              summary: "Model Registry Operator Probe Success Burn Rate"
              message: "High error budget burn for model-registry-operator (current value: 16)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-model-registry-operator-probe-success-burn-rate.md"

  - interval: 1m
    input_series: 
      - series: probe_success:burnrate2h{instance="model-registry-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="model-registry-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Model Registry Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Model Registry Operator Probe Success Burn Rate
              instance: "model-registry-operator"
              namespace: "redhat-ods-applications"
              severity: warning
            exp_annotations:
              summary: "Model Registry Operator Probe Success Burn Rate"
              message: "High error budget burn for model-registry-operator (current value: 61)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-model-registry-operator-probe-success-burn-rate.md"
